<a href="https://colab.research.google.com/github/PlaZMaD/GP_ML_2026/blob/main/Seminar_1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Добро пожаловать

На этом занятии:
1. Минимум Python: объекты, строки, контейнеры, функции, генераторы, классы
2. NumPy и векторизованные вычисления
3. Pandas: работа с табличными данными
4. Графики
5. Первые модели в scikit-learn: KNN и линейная регрессия
6. Переобучение

На практических занятиях курса мы будем использовать [язык программирования Python](https://www.python.org) в [среде Google Colab](https://colab.research.google.com). В качестве альтернативы вы можете загрузить какой-нибудь другой дистрибутив python, например [anaconda](https://www.anaconda.com/) и запустить jupyter локально (см. [docs](https://jupyter.readthedocs.io/en/latest/running.html)).

```python
!python --version
```

Если вы новичок в Python, пожалуйста, ознакомьтесь со следующим руководством:
 - https://docs.python.org/3/tutorial/

В частности, следующие части должны дать более или менее полное представление об основах, которые необходимо знать:
   - https://docs.python.org/3/tutorial/introduction.html
   - https://docs.python.org/3/tutorial/controlflow.html
   - https://docs.python.org/3/tutorial/datastructures.html
   - https://docs.python.org/3/tutorial/modules.html
   - https://docs.python.org/3/tutorial/classes.html

Не забывайте соблюдать [PEP-8](https://peps.python.org/pep-0008/). Вы также можете ознакомиться с другими [руководствами по стилю](https://google.github.io/styleguide/pyguide.html).

Обзор основных возможностей среды Google Colab можно найти в  [разделе](https://colab.research.google.com/notebooks/basic_features_overview.ipynb).

# Python: минимум для старта
## Объекты и переменные
В Python всё является объектом: числа, строки, списки, функции и даже классы. У каждого объекта есть тип (`type`) и набор методов — функций, которые вызываются через точку.

```python
x = 42
s = "hello"
print(type(x), type(s), type(print))

# методы — функции, «привязанные» к объекту
print(s.upper())
print(s.replace("l", "L"))

# переменная — это просто имя, которое ссылается на объект
a = [1, 2, 3]
b = a          # b ссылается на тот же самый список
b.append(4)
print(a)       # изменился и a!
print(a is b, a == [1, 2, 3, 4])
```

## Строки
Строки можно задавать в одинарных или двойных кавычках, а f-строки позволяют подставлять в текст значения переменных:

```python
a = 5
b = 6
c = a+b

my_string  = "sdfsdf'df'hgg"
my_string_2  = 'dsfs"d"fsdf'

my_string  = f"my good \n  {my_string} string"

print(my_string)
```

```python
text = "numpy,pandas,sklearn"
parts = text.split(",")        # разбить строку на список
print(parts)
print(" + ".join(parts))       # склеить список в строку
print(text[:5], text[-7:], len(text))  # срезы и длина
```

## Списки и словари

```python
my_list = [1, 2, 4, 5, 6]
my_list.append(7)
print(my_list)
```

```python
my_dict = {'key_1':'val_1'}
my_dict['key_2'] = 'val_2'
print(my_dict['key_2'])
```

## Функции
Функция объявляется с помощью `def`. У аргументов могут быть значения по умолчанию, а результат возвращается через `return`:

```python
def power(x, p=2):
    """Возводит x в степень p."""
    return x ** p

print(power(3))        # p по умолчанию равно 2
print(power(2, p=10))

# функции — тоже объекты: их можно передавать в другие функции
def apply_to_list(func, values):
    return [func(v) for v in values]

print(apply_to_list(power, [1, 2, 3]))
print(apply_to_list(lambda v: v + 100, [1, 2, 3]))  # lambda — короткая безымянная функция
```

## Генераторы и списковые включения
Некоторые объекты не хранят все элементы в памяти, а вычисляют их по одному по мере перебора. Например, `range(5)` не создаёт список из пяти чисел — элементы появляются во время итерации.
Перебирать элементы в цикле `for` можно у практически всех контейнеров, а списковые включения (list comprehension) позволяют коротко создавать новые контейнеры:

```python
a = range(5)
print(a)
for i in a:
  print (i)
```

```python
b = [i for i in a]
c = [i**2 for i in a if i%2==0]
d = {i:f"Квадратом числа {i} является {i**2}" for i in c}
print("b: ", b)
print("c: ", c)
print("d: ", d)
```

Генератор — это «ленивый» объект: он хранит только текущее состояние и правило получения следующего элемента. Создать его можно функцией с `yield` или выражением в круглых скобках:

```python
def squares(n):
    for i in range(n):
        yield i ** 2

gen = squares(5)
print(gen)                 # значения ещё не вычислены
print(next(gen), next(gen))
print(list(gen))           # оставшиеся элементы

print(sum(i ** 2 for i in range(5)))  # генераторное выражение
```

## Чуть-чуть о классах в Python
До этого мы работали только с готовыми типами объектов, но можно создавать и свои!

```python
class Dog:
  legs = 4
  def __init__(self, name, age):
      self.name = name
      self.age = age
  def speak(self):
    print("Ауф!")
```

```python
Nemo = Dog('Nemo', 6)
print(Nemo.name)
print(Nemo.age)
Nemo.speak()

Muhtar = Dog('Muhtar', 10)
print(Muhtar.name)
print(Muhtar.legs)
```

```python
# атрибут экземпляра «перекрывает» атрибут класса только у этого объекта
Muhtar.legs = 3
print(Muhtar.legs, Nemo.legs)

# изменение атрибута класса видно всем объектам, у которых нет своего значения
Dog.legs = 2
print(Muhtar.legs, Nemo.legs)
```

# NumPy и векторизованные вычисления

Почти любая модель машинного обучения требует заметных вычислений, обычно связанных с задачами линейной алгебры. К сожалению, «чистый» Python справляется с этим плохо: каждая операция по отдельности выполняется интерпретатором во время работы кода.

Поэтому мы воспользуемся numpy — библиотекой для очень быстрых вычислений с векторами, матрицами и другими тензорами. Основной объект здесь — `numpy.ndarray`:

**Запускайте все ячейки с кодом по ходу чтения!**

```python
import numpy as np

a = np.array([1,2,3,4,5])
b = np.array([5,4,3,2,1])
print("a = ", a)
print("b = ", b)

# math and boolean operations can applied to each element of an array
print("a + 1 =", a + 1)
print("a * 2 =", a * 2)
print("a == 2", a == 2)
# ... or corresponding elements of two (or more) arrays
print("a + b =", a + b)
print("a * b =", a * b)
```

<!-- grade_id: cell-c8488cc1b2d24a07, points: 0, grade: false, solution: true -->
```python
# Ваша очередь: вычислите поэлементные полупроизведения a и b (половины произведений)

def half_product(a, b):
    # ваш код здесь
    raise NotImplementedError
```

Убедитесь, что автоматические проверки пройдены!


<!-- grade_id: HalfProduct, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(
    half_product(
        np.array([ 1.3549439 , -1.69765972,  0.32111273, -0.2703243 , -0.38928596]),
        np.array([-1.50693732,  0.69443195,  0.90948151,  1.1717395 ,  1.4069652 ])
    ),
    np.array([-1.02090776, -0.58945457,  0.14602305, -0.15837483, -0.2738559 ])
)

print("a =", a)
print("b =", b)
print("a * b / 2 = ", half_product(a, b))
```

<!-- grade_id: cell-68ccf91f62d2ac06, points: 0, grade: false, solution: true -->
```python
# Вычислите поэлементное частное квадрата a и (b + 1), верните результат

def a_squared_over_b_plus_1(a, b):
    # ваш код здесь
    raise NotImplementedError
```

Убедитесь, что автоматические проверки пройдены!

<!-- grade_id: ASquaredOverBPlus1, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(
    a_squared_over_b_plus_1(
        np.array([ 1.3549439 , -1.69765972,  0.32111273, -0.2703243 , -0.38928596]),
        np.array([-1.50693732,  0.69443195,  0.90948151,  1.1717395 ,  1.4069652 ])
    ),
    np.array([-3.62149895,  1.70089364,  0.05400072,  0.03364825,  0.06296043])
)

print("a = ", a)
print("b = ", b)
print("a**2 / (b + 1)", a_squared_over_b_plus_1(a, b))
```

Существует множество функций для создания массивов нулей, единиц, возрастающих чисел и т.д.:

```python
np.zeros(shape=(3, 4))
```

```python
np.ones(shape=(2, 5))
```

```python
np.arange(3, 15, 2.5) # start, stop, step
```

```python
np.linspace(0, 10, 11) # divide [0, 10] interval into 11 points
```

Вы можете легко изменять форму массивов:

```python
np.arange(24).reshape(2, 3, 4)
```

При операциях над массивами разной формы размерности автоматически согласуются (broadcasting):

```python
a = np.arange(3).reshape(1, 3)
b = np.arange(3).reshape(3, 1)
print(a)
print(b)
a - b
```

Broadcasting можно использовать по-разному. Например, для поэлементного возведения матрицы в несколько степеней сразу:

```python
A = np.array([
    [1, 2],
    [3, 5]
])
powers = np.arange(1, 5)

A.reshape(1, 2, 2)**powers.reshape(4, 1, 1)
```

Существует также несколько способов объединения массивов. Например, `np.concatenate` объединяет массивы вдоль **существующей** оси:

```python
matrix1 = np.arange(18).reshape(6, 3) # array of shape (6, 3)
matrix2 = -np.arange(12).reshape(6, 2) # array of shape (6, 2)

np.concatenate([matrix1, matrix2], axis=1) # result of shape (6, 5)
```

А `np.stack` добавляет к результату новую ось:

```python
np.stack([
    np.arange(5), # array of shape (5,)
    np.linspace(0, 1, 5) # array of shape (5,)
], axis=1) # result of shape (5, 2)
```

Любая матрица может быть легко транспонирована:

```python
print(matrix2)
print('---')
print(matrix2.T)
```

```python
print('matrix2.shape =', matrix2.shape)
print('matrix2.T.shape =', matrix2.T.shape)
```

<!-- grade_id: cell-7052c430a2292e3f, points: 0, grade: false, solution: true -->
```python
# Ваша очередь!
# Используя изученные методы (создание массивов,
# изменение формы и объединение), напишите функцию, которая составляет
# и возвращает следующую матрицу:
#
#              0   7  14  21  28  0  0
#              1   8  15  22  29  0  0
#              2   9  16  23  30  0  0
#              3  10  17  24  31  3  3
#              4  11  18  25  32  3  3
#              5  12  19  26  33  3  3
#              6  13  20  27  34  3  3


def compose_matrix():
    # ваш код здесь
    raise NotImplementedError
```

<!-- grade_id: ComposeMatrix, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal(
    compose_matrix(),
    np.array([[0,   7,  14,  21,  28,  0,  0],
              [1,   8,  15,  22,  29,  0,  0],
              [2,   9,  16,  23,  30,  0,  0],
              [3,  10,  17,  24,  31,  3,  3],
              [4,  11,  18,  25,  32,  3,  3],
              [5,  12,  19,  26,  33,  3,  3],
              [6,  13,  20,  27,  34,  3,  3]])
)
```

Линейная алгебра:

```python
a = np.arange(3)
B = np.arange(12).reshape(4, 3)
print("a =", a)
print("B =\n", B)
print("Dot product (a*a):", a @ a) # or: np.dot(a, a)
print("Matrix-vector (B*a):", B @ a) # or: np.matmul(B, a)
print("Matrix-matrix (B*B.T):\n", B @ B.T) # or: np.matmul(B, B.T)
```

И огромное количество прочих математических операций:

```python
a = np.array([1,2,3,4,5])
b = np.array([5,4,3,2,1])
print("numpy.sum(a) = ", np.sum(a))
print("numpy.mean(a) = ", np.mean(a))
print("numpy.min(a) = ",  np.min(a))
print("numpy.argmin(b) = ", np.argmin(b))  # index of minimal element
print("numpy.dot(a,b) = ", np.dot(a, b))      # dot product. Also used for matrix/tensor multiplication
print("numpy.unique(['male','male','female','female','male']) = ", np.unique(['male','male','female','female','male']))

# and tons of other stuff. see https://numpy.org/doc/stable/reference/routines.html
```

```python
# most of this functions are also implemented as members of numpy arrays, e.g.:
print('a.min() =', a.min())
print('a.mean() =', a.mean())
```

Насколько numpy быстрее «чистого» Python? Сравним поэлементное сложение двух массивов из 100 000 чисел:

```python
X_list = list(range(100000))
Y_list = list(range(100000))
X = np.array(X_list)
Y = np.array(Y_list)

def pure_python_version():
    Z = []
    for i in range(len(X_list)):
        Z.append(X_list[i] + Y_list[i])

def numpy_version():
    Z = X + Y
```

```python
%timeit pure_python_version()
```

```python
%timeit numpy_version()
```

# Pandas -  работа с табличными данными

Pandas — это библиотека, которая поможет загрузить данные, подготовить их и провести несложный анализ. Основной объект здесь — `pandas.DataFrame`, двумерная таблица (под капотом работает numpy).

Загрузим наш датасет:

```python
!wget https://raw.githubusercontent.com/HSE-LAMBDA/MLDM-2022/main/01-intro/tabular-playground-series-aug-2022.zip
```

```python
!unzip -o tabular-playground-series-aug-2022.zip -d data
```

Теперь прочитаем данные с помощью pandas:

```python
import pandas as pd
data = pd.read_csv('data/train.csv',
                    index_col='id') # this yields a pandas.DataFrame
data.index += 1
```

```python
# Take a look at the data

data.head(5) # selects top 5 lines
```

```python
pd.options.display.max_rows = 10
data
```

```python
data.sample(5)
```

#### О данных
Эти данные — результаты крупного исследования по тестированию продукции. Для каждого кода продукта (`product_code`) дан набор атрибутов (`attribute_*`), общих для всех образцов с этим кодом, а также набор измерений (`measurement_*`) для каждого отдельного образца, полученных разными методами лабораторных испытаний. Каждый образец испытывается в условиях, имитирующих реальную эксплуатацию: он поглощает определённое количество жидкости (`loading`), и проверяется, не выйдет ли он из строя.

Задача — по результатам лабораторных испытаний предсказать выход из строя (`failure`) для образцов с новыми кодами продукта. Обучающие данные содержат целевую переменную `failure`; нужно предсказать вероятность того, что образец выйдет из строя.

```python
# table dimensions
print("len(data) = ", len(data.index))
print("data.shape = ", data.shape)
```

```python
# select a single row
print(data.loc[4])
```

```python
# select a single column.
loadings = data["loading"] # alternatively: data.loading
print(loadings.loc[:10])
```

```python
# select several columns and rows at once
data.loc[5:10, ("loading", "failure")]    # alternatively: data[["loading","failure"]].loc[5:10]
```

### `loc` vs `iloc`

В pandas есть два способа индексирования строк:

*  по значениям столбцов и индекса (в нашем случае id) - для этого используется data.loc
*  по позиционному индексу - для этого используется data.iloc

Обратите внимание, что индексный столбец начинается с 1, поэтому позиционный индекс 0 будет соответствовать значению индексного столбца 1, позиционный 1 - значению индексного столбца 2, и так далее:

```python
print(data.index)
print('------')
print("data.iloc[0]:")
print(data.iloc[0])
print('------')
print("data.loc[1]:")
print(data.loc[1])
```

Также обратите внимание, что при индексировании с помощью `.loc` включаются оба конца среза:

```python
data.loc[2:3]
```

в то время как при использовании `.iloc` конец исключается:

```python
data.iloc[1:2]
```

Более сложная индексация (похожая на булевую индексацию в numpy):

```python
data.loc[(data['loading'] < 90) & (data['product_code'] == np.random.choice(data.product_code.unique()))]
```

```python
data.query('2.1 * measurement_9 < measurement_10 and loading < 111')
```

### Ваша очередь:

```python
# Выберите исследования с номерами 13 и 666 одной строкой — вышли ли образцы из строя?

print(data.loc[[13, 666], 'failure'])
```

```python
# Вычислите общую долю отказов (какая доля образцов вышла из строя).
# Сбалансирована ли обучающая выборка?

print(data['failure'].mean())
```

В Pandas также есть некоторые базовые инструменты анализа данных. Например, вы можете быстро вывести статистические показатели для каждого столбца с помощью `.describe()`.

```python
data.describe()
```

Некоторые столбцы содержат значения __NaN__ - это означает, что в них нет данных. Например, у образца с `id` 10 неизвестно значение `measurement_16`. Чтобы упростить дальнейший анализ данных, мы заменим значения NaN с помощью функции pandas `fillna`.

Примечание: мы делаем это, потому что это учебный материал. В общем случае, подумайте дважды, прежде чем изменять данные подобным образом.

```python
data['measurement_9'] = data['measurement_9'].fillna(value=data['measurement_9'].mean())
data['measurement_16'] = data['measurement_16'].fillna(value=data['measurement_16'].mean())
```

```python
data.loc[10]
```

### Pandas + numpy

Важно: поскольку pandas построен поверх numpy, большинство функций numpy работает и с датафреймами. Представление данных в виде numpy-массива можно получить через `.values` (или `.to_numpy()`), но многие функции numpy работают и напрямую с объектами pandas:

```python
# calling np.max on a pure pandas column:
column_name = 'measurement_17'
print("Max {}: ".format(column_name), np.max(data[column_name]))

# calling np.argmax on a numpy representation of a pandas column
# to get its positional index:
print("\nThe study with the max " + column_name + ":\n",
      data.iloc[
          np.argmax(data[column_name].values)
      ])
```

```python
# numpy works only with positional index:
column_name = 'measurement_16'
print(data[column_name].values.argmax())
#     ^^^^^^^^^^^^^^^^^^^
#     this part returns a numpy array, argmax of which we are calculating


# in pandas you can ask for the index (i.e. value of the index column)
# of the maximal element like this:
print(data[column_name].idxmax())
```

### Ваша очередь

Используя numpy и pandas, ответьте на несколько вопросов о данных.

```python
# Ваш код: вычислите среднее значение loading и найдите максимальное.

# ВАШ КОД
```

```python
# Образцы с каким кодом продукта чаще выходят из строя?
prod_names = data.product_code.unique()
print(prod_names)
# ВАШ КОД
```

# Графики

Самые популярные библиотеки для построения графиков — Matplotlib (статичные графики) и Plotly (интерактивные).

```python
import matplotlib.pyplot as plt
%matplotlib inline

import plotly.graph_objects as go
```

```python
n_points = 5
x = np.arange(n_points)
y = x**2

print("x =", x)
print("x**2 =", y)
print("plotting x**2 vs x:")


plt.scatter(x, y)
plt.show()  # show the first plot to begin drawing the next one

fig = go.Figure()
fig.add_trace(go.Scatter(x=x, y=y, mode='markers'))
fig.show()
```

```python
# histogram - showing data density
float_cols = [c for c in data.columns if data[c].dtype == float]

_, axs = plt.subplots(2, 4, figsize=(10,5))

for f, ax in zip(float_cols[:8], axs.ravel()):
    bins = np.linspace(data[f].min(), data[f].max(), 50)
    ax.hist(data[f], bins=bins, density=True)
    ax.set_xlabel(f)

plt.tight_layout(w_pad=1)
plt.suptitle('Distributions of the continuous features', fontsize=15, y=1.12)
plt.show()
```

```python
# .plot() method allows you to plot in different styles
data.product_code.value_counts().plot(kind='pie', autopct="%1.1f%%",shadow=True,
        startangle=45, explode=[0.065] * data.product_code.nunique());
```

# Машинное обучение со scikit-learn

Scikit-learn — это библиотека для построения простых конвейеров (pipeline) машинного обучения.

Она объединяет множество моделей под общим интерфейсом:

- Создание модели: `model = sklearn.whatever.ModelNameHere(parameters_if_any)`
- Обучение: `model.fit(X, y)`
- Предсказание: `model.predict(X_test)`

В ней также есть инструменты для извлечения признаков, оценки качества, кросс-валидации и многого другого.

# Метод k ближайших соседей (K-nearest neighbors)

```python
from sklearn.neighbors import KNeighborsClassifier
from sklearn.metrics import accuracy_score
from sklearn.model_selection import train_test_split


x = data[float_cols].fillna(0).copy()
y = data["failure"]

model = KNeighborsClassifier(n_neighbors=5)

# split the data into train(90%) and test(10%)
X_train, X_test, y_train, y_test = train_test_split( x, y, test_size=0.1, random_state=42)


# fit the model
model.fit(X_train,  y_train)

# make the prediction
test_predictions = model.predict(X_test)
print("Test accuracy:", accuracy_score(y_test, test_predictions))

```

Много это или мало? Сравним с простейшим «классификатором», который всегда предсказывает самый частый класс (`failure = 0`):

```python
from sklearn.dummy import DummyClassifier

baseline = DummyClassifier(strategy="most_frequent")
baseline.fit(X_train, y_train)
print("Baseline accuracy:", accuracy_score(y_test, baseline.predict(X_test)))
```

KNN оказался **хуже константы**! Всегда сравнивайте модель с простым базовым решением (baseline).

Одна из причин: KNN считает расстояния между объектами, поэтому признаки с большим разбросом значений (например, `loading`) «перевешивают» остальные. Перед KNN признаки нужно масштабировать. Объединим масштабирование и модель в цепочку (pipeline):

```python
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import StandardScaler

model = make_pipeline(StandardScaler(), KNeighborsClassifier(n_neighbors=5))
model.fit(X_train, y_train)
print("Test accuracy (with scaling):", accuracy_score(y_test, model.predict(X_test)))
```

Масштабирование — обязательный шаг для KNN, но здесь оно не помогло: константу мы так и не обогнали. Дело в том, что классы несбалансированы (≈79 % образцов не выходят из строя), а по этим признакам отказ предсказать трудно. Accuracy в такой ситуации — плохая метрика; о метриках качества поговорим на следующем занятии.

# Линейная регрессия

## y = kx + b

Начнём с одномерного случая, где истинная зависимость имеет вид:
$$y=k\cdot x+b+\text{noise}$$

```python
def linear_function(x):
  return 0.33 * x + 8.3

def gen_dataset(N, func, lims=(-1., 1.), noise_lvl=0.2):
  x = np.random.uniform(*lims, size=N)
  y = func(x) + noise_lvl * np.random.normal(size=x.shape)
  return x[:,None], y

np.random.seed(42)
X, y = gen_dataset(50, linear_function)
x = np.linspace(-1, 1, 101)
plt.plot(x, linear_function(x))
plt.scatter(X, y);
```

```python
from sklearn.linear_model import LinearRegression

model = LinearRegression()

model.fit(X, y)

x = np.linspace(-1, 1, 101)
plt.plot(x, linear_function(x), label='true function')
plt.scatter(X, y);
plt.plot(x, model.predict(x[:,None]), label='prediction')
plt.legend()

print(model.coef_, model.intercept_)
```

А что если случай более сложный?

```python
def true_function(x):
  return np.sin(3 * x + 0.8) + np.sin(1. / (x + 1.23))
```

```python
model = LinearRegression()
X, y = gen_dataset(25, true_function)

model.fit(X, y)

x = np.linspace(-1, 1, 101)
plt.plot(x, true_function(x), label='true function')
plt.scatter(X, y);
plt.plot(x, model.predict(x[:,None]), label='prediction')
plt.legend();
```

На самом деле даже довольно сложные задачи можно решать простыми алгоритмами, если хорошо понимать саму задачу!
В данном случае поможет конструирование новых признаков (фичей) — например, степеней исходного признака $x$.
Для удобства цепочку обработки данных мы будем помещать в pipeline.


$$X'=
\begin{pmatrix}
x_1 & (x_1)^2 & \ldots & (x_1)^k \\
x_2 & (x_2)^2 & \ldots & (x_2)^k \\
\vdots & \vdots & \ddots & \vdots \\
x_N & (x_N)^2 & \ldots & (x_N)^k
\end{pmatrix},
$$

```python
from sklearn.preprocessing import PolynomialFeatures
from sklearn.pipeline import make_pipeline

poly_expand = PolynomialFeatures(3)
poly_expand.fit_transform(np.arange(5)[:,None])
```

```python
# Первый параметр — степень полинома. Попробуйте разные значения: 1, 3, 5, 15, 50
poly_expand = PolynomialFeatures(5, include_bias=False)
linear_model = LinearRegression()
model = make_pipeline(
    poly_expand, linear_model
)

model.fit(X, y)

x = np.linspace(-1, 1, 101)
plt.plot(x, true_function(x), label='true function')
plt.scatter(X, y);
plt.plot(x, model.predict(x[:,None]), label='prediction')
plt.ylim(y.min() - 0.5, y.max() + 0.5)
plt.legend();
```

# Рассмотрим более сложный датасет - Boston Housing

Описание переменных:

- **CRIM** — уровень преступности на душу населения по городу
- **ZN** — доля жилой земли под участки площадью более 25 000 кв. футов
- **INDUS** — доля площадей, занятых нерозничным бизнесом, по городу
- **CHAS** — бинарная переменная «река Чарльз» (= 1, если участок граничит с рекой; 0 иначе)
- **NOX** — концентрация оксидов азота (частей на 10 млн)
- **RM** — среднее число комнат в жилище
- **AGE** — доля домов, занимаемых владельцами и построенных до 1940 г.
- **DIS** — взвешенное расстояние до пяти деловых центров Бостона
- **RAD** — индекс доступности радиальных магистралей
- **TAX** — ставка налога на недвижимость на каждые \$10 000
- **PTRATIO** — соотношение числа учеников и учителей по городу
- **B** — $1000(B_k - 0.63)^2$, где $B_k$ — доля чернокожего населения по городу
- **LSTAT** — доля населения с низким социальным статусом, %
- **MEDV** — медианная стоимость домов, занимаемых владельцами, в тысячах \$

```python
!wget https://raw.githubusercontent.com/selva86/datasets/master/BostonHousing.csv
```

```python
raw_df = pd.read_csv('BostonHousing.csv')
print(raw_df.head())

y = raw_df['medv']
X = raw_df.drop('medv', axis=1)
```

Мы хотим предсказать MEDV - цену жилья, исходя из остальных показателей. Оценим зависимость целевой переменной от фичей:

```python
plt.figure(figsize=(15, 15))

grid_size = int(np.ceil(X.shape[1]**0.5))

for i, (name, x) in enumerate(X.items(), 1):
  plt.subplot(grid_size, grid_size, i)
  plt.scatter(x, y)
  plt.xlabel(name)
```

```python
from sklearn.linear_model import LinearRegression
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import PolynomialFeatures
from sklearn.metrics import mean_squared_error
```

Выберем 3 наиболее удачных признака и построим модель на их основе.

```python
old_columns = X.columns
print(old_columns)
X.columns = [column.upper() for column in X.columns]
```

```python
columns = ["CRIM", "RM", "LSTAT"]

X_subset = X[columns].copy()
X_subset /= X_subset.max()
```

```python
model = make_pipeline(
    PolynomialFeatures(9, include_bias=False), # Can you calculate how many features this will result in? :)
    LinearRegression()
)

model.fit(X_subset, y)
print('mse = ', mean_squared_error(y, model.predict(X_subset)))

plt.figure(figsize=(15, 4))
for i, c in enumerate(columns, 1):
  plt.subplot(1, len(columns), i)
  plt.scatter(X[c], y, label='data')
  plt.scatter(X[c], model.predict(X_subset), label='prediction')
  plt.legend()
```

Насколько это хорошо? На вид здорово, а что на самом деле?

# Разбиение на обучающую и тестовую выборки (train-test split)

```python
X_train, X_test, y_train, y_test = train_test_split(X_subset, y, test_size=50, random_state=39)

model = make_pipeline(
    PolynomialFeatures(9, include_bias=False),
    LinearRegression()
)

model.fit(X_train, y_train)

print('train mse = ', mean_squared_error(y_train, model.predict(X_train)))
print('test mse = ', mean_squared_error(y_test, model.predict(X_test)))
```

Что?! Ошибка на тестовой выборке намного больше, чем на обучающей.

```python
plt.figure(figsize=(15, 4))
for i, c in enumerate(columns, 1):
  plt.subplot(1, len(columns), i)
  plt.scatter(X_test[c], y_test, label='data')
  plt.scatter(X_test[c], model.predict(X_test), label='prediction')
  plt.legend()
```

Может что-то не то с коэффициентами?

```python
model.named_steps['linearregression'].coef_
```

## Итог
На обучающей выборке модель выглядит отлично, а на новых данных ошибается гораздо сильнее — это **переобучение** (overfitting). Полином 9-й степени от 3 признаков даёт 219 признаков, и огромные значения коэффициентов — характерный симптом.

Как с этим бороться, разберём на следующем занятии.